---
title: 'Introduction to Programming'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: a7d1da812c
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}
print ("aa")
```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
